<a href="https://colab.research.google.com/github/fedeguzani/simulador/blob/main/tp2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Mi primera red neuronal
En este practico trabajaremos con la base de datos de MNIST. Esta base de datos consta de digitos escritos a mano. Para el final de este trabajo practico, habran implementado una red neuronal profunda que sepa reconocer cada uno de estos digitos.

### Atencion! Trata de no cambiar los nombres de las variables! Recuerda que testearemos estas pruebas automaticamente.

## Pregunta 1
### Importar los datos
Lo primero que tendras que hacer es bajar la base de datos. Por suerte, esta no es una tarea demasiado dificil, ya que el mismo tensorflow la posee y podras buscar en su codigo como cargarla a memoria.
Sin embargo, tienes que investigar como hacerlo.
[datasets de tensorflow](https://www.tensorflow.org/api_docs/python/tf/keras/datasets)

```python
import tensorflow as tf
from tensorflow import keras
```

<!-- grade_id: Primera_pregunta, points: 0, grade: false, solution: true -->
```python
# Descomenta (saca la #) la proxima linea de codigo e implementa la linea
# que bajara la base de datos mnist
 (x_train, train_labels), (x_test, test_labels) = # linea_de_codigo_aqui


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: Primera_prueba, points: 1, grade: true, solution: false -->
```python
# TEST CELL
assert x_train is not None, "No has cargado los datos de entrenamiento"
assert x_test is not None, "No has cargado los datos de prueba"

```

```python
# Normalizando los datos
x_train, x_test = x_train / 255.0, x_test / 255.0
```

## Pregunta 2
### Parte 2: Creación del Modelo

Crea un modelo de red neuronal utilizando la API de Keras en TensorFlow. Tu modelo debería seguir la estructura siguiente:

1. Una capa de aplanamiento (`Flatten`) que convierta las imágenes de 28x28 píxeles en un vector unidimensional.
2. Una capa densa (`Dense`) con 128 neuronas y función de activación ReLU.
3. Una capa de dropout (`Dropout`) con una tasa de 0.2, para ayudar a prevenir el sobreajuste.
4. Una capa densa final (`Dense`) con 10 neuronas (una para cada dígito del 0 al 9) y función de activación softmax (para la clasificación multiclase).

Guarda tu modelo en una variable llamada `model`.


<!-- grade_id: Segunda_pregunta, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: Segunda_prueba, points: 2, grade: true, solution: false -->
```python
# TEST CELL
assert model is not None, "No has creado el modelo!"
assert isinstance(model, tf.keras.models.Sequential), "Tu modelo debe ser una instancia de tf.keras.models.Sequential"
assert len(model.layers) == 4, "Tu modelo debe tener exactamente 4 capas"

```

```python

Ahora compila tu modelo utilizando el optimizador `adam` y la función de pérdida `sparse_categorical_crossentropy`. También debes monitorear la métrica de `accuracy`.

```

<!-- grade_id: Tercera_pregunta, points: 0, grade: false, solution: true -->
```python


# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: Tercera_prueba, points: 2, grade: true, solution: false -->
```python
# TEST CELL
assert model.optimizer is not None, "No has configurado un optimizador para tu modelo"
assert model.loss is not None, "No has configurado una función de pérdida para tu modelo"
assert "accuracy" in model.metrics_names, "Debes monitorear la métrica 'accuracy'"

```

## Pregunta 3
### Parte 3: Entrenamiento del Modelo

Entrena tu modelo en los datos de entrenamiento durante al menos 5 épocas y guarda el objeto de historial devuelto por `model.fit()` en una variable llamada `history`. No olvides también proporcionar los datos de validación.



<!-- grade_id: Cuarta_pregunta, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: Cuarto_test, points: 3, grade: true, solution: false -->
```python
# TEST CELL
assert history is not None, "No has entrenado el modelo"
assert len(history.history['loss']) >= 5, "Debes entrenar el modelo durante al menos 5 épocas"


```

```python
Una vez que tu modelo esté entrenado, visualiza la evolución de la precisión de entrenamiento y validación a lo largo de las épocas. Puedes usar `matplotlib.pyplot` para esto.

```

<!-- grade_id: Visualizacion, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()

```

Aquí te mostramos un ejemplo de cómo debería verse tu gráfico. La precisión de entrenamiento y validación deberían aumentar con el tiempo, y la precisión de validación debería ser algo menor que la de entrenamiento.

Recuerda, la precisión exacta puede variar debido a la naturaleza estocástica del entrenamiento de la red neuronal, por lo que no te preocupes si tus números no coinciden exactamente con los del ejemplo.


<!-- grade_id: Ejemplo_visualizacion_nump, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Pregunta 4
### Parte 4: Evaluación del Modelo

Ahora que has entrenado tu modelo, debes evaluar su rendimiento en los datos de prueba. Usa la función `model.evaluate()` para hacerlo y guarda el resultado en una variable llamada `test_loss, test_acc`.


<!-- grade_id: Pregunta_evaluacion, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: Prueba_evaluacion, points: 2, grade: true, solution: false -->
```python
# TEST CELL
assert test_acc is not None, "No has evaluado el modelo en los datos de prueba"
assert 0 <= test_acc <= 1, "La precisión debe estar entre 0 y 1"

```

Finalmente, imprime la precisión del modelo en los datos de prueba.


<!-- grade_id: Imprime_precision, points: 0, grade: true, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()

```

## Pregunta 5
### Parte 5: Interpretación de Resultados

Por favor, escribe un breve párrafo en el que interpretes los resultados obtenidos. ¿Cómo evolucionó la precisión de tu modelo durante el entrenamiento? ¿Cómo se compara la precisión de entrenamiento con la de validación? ¿Estás satisfecho con la precisión del modelo en los datos de prueba?


<!-- grade_id: Respuesta_teorica, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE